Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Daniel Wossen"
import pandas as pd
import numpy as np
df = pd.read_csv('Salaries.csv')
df.head()
```

Output:
```
C:\Users\hp\AppData\Local\Temp\ipykernel_20972\1375400659.py:4: DtypeWarning: Columns (3,4,5,6,12) have mixed types. Specify dtype option on import or set low_memory=False.
  df = pd.read_csv('Salaries.csv')
   Id       EmployeeName                                        JobTitle  \
0   1     NATHANIEL FORD  GENERAL MANAGER-METROPOLITAN TRANSIT AUTHORITY   
1   2       GARY JIMENEZ                 CAPTAIN III (POLICE DEPARTMENT)   
2   3     ALBERT PARDINI                 CAPTAIN III (POLICE DEPARTMENT)   
3   4  CHRISTOPHER CHONG            WIRE ROPE CABLE MAINTENANCE MECHANIC   
4   5    PATRICK GARDNER    DEPUTY CHIEF OF DEPARTMENT,(FIRE DEPARTMENT)   

     BasePay OvertimePay   OtherPay Benefits   TotalPay  TotalPayBenefits  \
0  167411.18         0.0  400184.25      NaN  567595.43         567595.43   
1  155966.02   245131.88  137811.38      NaN  538909.28         538909.28   
2  212739.13   106088.18    16452.6      NaN  335279.91         335279.91   
3    77916.0    56120.71   
```

---

# Q3: Pandas

Welcome to a quick exercise for you to practice your pandas skills! We will be using the [SF Salaries Dataset](https://www.kaggle.com/kaggle/sf-salaries) from Kaggle! Just follow along and complete the tasks outlined in bold below. The tasks will get harder and harder as you go along.

Write a function `average_base_pay(df)` that takes a dataframe and returns the average BasePay.


<!-- grade_id: cell-9956d2606b32c334, points: 0, grade: false, solution: true -->
```python

def average_base_pay(df):
    """Compute the average base pay in the dataframe.
    
    Args:
        df: A dataframe with a column 'BasePay' containing numerical values.
    
    Returns:
        The average base pay as a float.
    """
    
    base_pay_col = df['BasePay']
    base_pay_col = base_pay_col.replace('Not Provided', np.nan)
    base_pay_col = base_pay_col.dropna()
    base_pay_col = base_pay_col.astype(float)
    return base_pay_col.mean()
    


```

<!-- grade_id: cell-6de9cf25de12ebfe, points: 0, grade: true, solution: false -->
```python
def test_average_base_pay():
    """Check that average_base_pay returns the correct output for several inputs"""
    assert average_base_pay(df) == 66325.44884048768
    print("All tests passed")

test_average_base_pay()

```

Output:
```
All tests passed
```

q) Write a function `highest_overtime_pay(df)` that takes a dataframe and returns the highest OvertimePay.


<!-- grade_id: cell-a21f0da23bd597e0, points: 0, grade: false, solution: true -->
```python
def highest_overtime_pay(df):
    """Compute the highest overtime pay in the dataframe.
    
    Args:
        df: A dataframe with a column 'OvertimePay' containing numerical values.
    
    Returns:
        The highest overtime pay as a float.
    """
    over_time_col = df['OvertimePay']
    over_time_col = over_time_col.replace('Not Provided', np.nan)
    over_time_col = over_time_col.dropna()
    over_time_col = over_time_col.astype(float)
    return over_time_col.max()

```

<!-- grade_id: cell-0b7a1b31855d854b, points: 0, grade: true, solution: false -->
```python
def test_highest_overtime_pay():
    """Check that highest_overtime_pay returns the correct output for several inputs"""
    print(highest_overtime_pay(df))

test_highest_overtime_pay()
```

Output:
```
245131.88
```

Write a function `job_title_of(df, name)` that takes a dataframe and a name, and returns the job title of the person with the given name. Make sure the name is case insensitive.

<!-- grade_id: cell-259ef69f6cbe9d0d, points: 0, grade: false, solution: true -->
```python
def job_title_of(df, name):
    """Find the job title of a person with a given name in the dataframe.
    
    Args:
        df: A dataframe with columns 'EmployeeName' and 'JobTitle' containing strings.
        name: A string representing a name.
    
    Returns:
        The job title as a string.
    """
    person_row = df[df['EmployeeName'] == name]
    # print(person_row['JobTitle'].values)
    job_title = person_row['JobTitle'].values[0]
    return job_title

```

<!-- grade_id: cell-dccc5ee19e3b6f65, points: 0, grade: true, solution: false -->
```python
def test_job_title_of():
    """Check that job_title_of returns the correct output for several inputs"""
    print(job_title_of(df, "GARY JIMENEZ"))

test_job_title_of()
```

Output:
```
['CAPTAIN III (POLICE DEPARTMENT)']
CAPTAIN III (POLICE DEPARTMENT)
```

Write a function `total_pay_benefits_of(df, name)` that takes a dataframe and a name, and returns the total pay benefits of the person with the given name. Make sure the name is case insensitive.


<!-- grade_id: cell-921b58b6ee68ca5c, points: 0, grade: false, solution: true -->
```python
def total_pay_benefits_of(df, name):
    """Find the total pay benefits of a person with a given name in the dataframe.
    
    Args:
        df: A dataframe with columns 'EmployeeName' and 'TotalPayBenefits' containing strings and numerical values, respectively.
        name: A string representing a name.
    
    Returns:
        The total pay benefits as a float.
    """
    # YOUR CODE HERE
    return df[df['EmployeeName'] == name]['TotalPayBenefits'].values[0]
```

<!-- grade_id: cell-e642664ad758145e, points: 0, grade: true, solution: false -->
```python
def test_total_pay_benefits_of():
    """Check that total_pay_benefits_of returns the correct output for several inputs"""
    
    

```

Write a function `highest_paid_person(df)` that takes a dataframe and returns the name of the highest paid person (including benefits).


<!-- grade_id: cell-e09c2726829088ad, points: 0, grade: false, solution: true -->
```python
def highest_paid_person(df):
    """Find the name of the highest paid person in the dataframe.
    
    Args:
        df: A dataframe with columns 'EmployeeName' and 'TotalPayBenefits' containing strings and numerical values, respectively.
    
    Returns:
        The name of the highest paid person as a string.
    """
    return df.loc[df['TotalPayBenefits'].idxmax()]['EmployeeName']
    
    

```

<!-- grade_id: cell-76533f6a7b16ae5b, points: 0, grade: true, solution: false -->
```python
def test_highest_paid_person():
    """Check that highest_paid_person returns the correct output for several inputs"""

```

Write a function `lowest_paid_person(df)` that takes a dataframe and returns the name of the lowest paid person (including benefits).


<!-- grade_id: cell-f3a8d3499aacc89d, points: 0, grade: false, solution: true -->
```python
def lowest_paid_person(df):
    """Find the name of the lowest paid person in the dataframe.
    
    Args:
        df: A dataframe with columns 'EmployeeName' and 'TotalPayBenefits' containing strings and numerical values, respectively.
    
    Returns:
        The name of the lowest paid person as a string.
    """
    
    # YOUR CODE HERE
    return df.loc[df['TotalPayBenefits'].idxmin()]['EmployeeName']

```

<!-- grade_id: cell-1b2727db36c9989c, points: 0, grade: true, solution: false -->
```python
def test_lowest_paid_person():
    """Check that lowest_paid_person returns the correct output for several inputs"""

```

Write a function `average_base_pay_per_year(df, year_start, year_end)` that takes a dataframe and a start and end year, and returns the average (mean) BasePay of all employees per year within the given range (inclusive).


<!-- grade_id: cell-acd1163d0716dd09, points: 0, grade: false, solution: true -->
```python
def average_base_pay_per_year(df, year_start, year_end):
    """Compute the average base pay per year in the given range.
    
    Args:
        df: A dataframe with columns 'BasePay' and 'Year' containing numerical values.
        year_start: An integer representing the start year.
        year_end: An integer representing the end year.
    
    Returns:
        A dictionary where the keys are the years and the values are the average base pay for that year.
    """
    # YOUR CODE HERE
    # raise NotImplementedError()
    df = df[(df['Year'] >= year_start) & (df['Year'] <= year_end)]
    df = df.groupby('Year')['BasePay'].mean()
    return df.to_dict()



```

<!-- grade_id: cell-5284dcb443a3ecd5, points: 0, grade: true, solution: false -->
```python
def test_average_base_pay_per_year():
    """Check that average_base_pay_per_year returns the correct output for several inputs"""
average_base_pay_per_year(df, 2011, 2012)
```

Output:
```
{2011: 63595.956516774524, 2012: 65436.40685742255}
{2011: 63595.956516774524, 2012: 65436.40685742255}
```

Write a function `unique_job_titles(df)` that takes a dataframe and returns the number of unique job titles.


<!-- grade_id: cell-6b9e50a1d9a2e2d8, points: 0, grade: false, solution: true -->
```python
def unique_job_titles(df):
    """Find the number of unique job titles in the dataframe.
    
    Args:
        df: A dataframe with a column 'JobTitle' containing strings.
    
    Returns:
        The number of unique job titles as an integer.
    """
    # # YOUR CODE HERE
    # raise NotImplementedError()
    return df['JobTitle'].nunique()

```

<!-- grade_id: cell-91fb07925e0f087c, points: 0, grade: true, solution: false -->
```python
def test_unique_job_titles():
    """Check that unique_job_titles returns the correct output for several inputs"""
```

Write a function `top_5_common_jobs(df)` that takes a dataframe and returns the top 5 most common jobs.


<!-- grade_id: cell-26edb4365b3fd49e, points: 0, grade: false, solution: true -->
```python
def top_5_common_jobs(df):
    """Find the top 5 most common jobs in the dataframe.
    
    Args:
        df: A dataframe with a column 'JobTitle' containing strings.
    
    Returns:
        A list of strings representing the top 5 most common jobs.
    """
    # YOUR CODE HERE
    # raise NotImplementedError()
    return df['JobTitle'].value_counts().head(5).index.tolist()

```

<!-- grade_id: cell-016df5e14047b530, points: 0, grade: true, solution: false -->
```python
def test_top_5_common_jobs():
    """Check that top_5_common_jobs returns the correct output for several inputs"""
print("Top 5 common jobs: ", top_5_common_jobs(df))
```

Output:
```
Top 5 common jobs:  ['Transit Operator', 'Special Nurse', 'Registered Nurse', 'Public Svc Aide-Public Works', 'Police Officer 3']
```

Write a function `job_titles_with_one_occurence(df, year)` that takes a dataframe and a year, and returns the number of job titles that were represented by only one person in that year.


<!-- grade_id: cell-aac7975288e0a390, points: 0, grade: false, solution: true -->
```python
def job_titles_with_one_occurence(df, year):
    """Find the number of job titles that were represented by only one person in a given year.
    
    Args:
        df: A dataframe with columns 'JobTitle' and 'Year' containing strings and numerical values, respectively.
        year: An integer representing a year.
    
    Returns:
        The number of job titles represented by only one person as an integer.
    """
    # YOUR CODE HERE
    # raise NotImplementedError()
    df_year = df[df['Year']==year]
    df_year = df_year.groupby('JobTitle').count()
    df_year = df_year[df_year['Year']==1]
    print(df_year)
    return df_year.shape[0]

```

<!-- grade_id: cell-c96c7770776e7a41, points: 0, grade: true, solution: false -->
```python
def test_job_titles_with_one_occurence():
    """Check that job_titles_with_one_occurence returns the correct output for several inputs"""
# print(job_titles_with_one_occurence(df, 2011))
```

Write a function `employees_with_chief_in_title(df)` that takes a dataframe and returns the number of people who have the word "Chief" in their job title. Make sure to account for different cases and placements of the word "Chief".


<!-- grade_id: cell-7308281b10b89d18, points: 0, grade: false, solution: true -->
```python
def employees_with_chief_in_title(df):
    """Find the number of people who have the word "Chief" in their job title.
    
    Args:
        df: A dataframe with a column 'JobTitle' containing strings.
    
    Returns:
        The number of people who have "Chief" in their job title as an integer.
    """
    # YOUR CODE HERE
    # raise NotImplementedError()
    return df[df['JobTitle'].str.contains('Chief')].shape[0]

```

<!-- grade_id: cell-8285b455e8aa75a1, points: 0, grade: true, solution: false -->
```python
def test_employees_with_chief_in_title():
    """Check that employees_with_chief_in_title returns the correct output for several inputs"""

```

Write a function `correlation_title_length_salary(df)` that takes a dataframe and returns the correlation between length of the Job Title string and Salary.

<!-- grade_id: cell-59b173bfaba1f9ab, points: 0, grade: false, solution: true -->
```python
def correlation_title_length_salary(df):
    """Compute the correlation between length of the job title and salary.
    
    Args:
        df: A dataframe with columns 'JobTitle' and 'TotalPayBenefits' containing strings and numerical values, respectively.
    
    Returns:
        The correlation as a float.
    """
    # YOUR CODE HERE
    # raise NotImplementedError()
    df['TitleLength'] = df['JobTitle'].apply(len)
    return df['TitleLength'].corr(df['TotalPayBenefits'])

```

<!-- grade_id: cell-0f20891084c2d2ed, points: 0, grade: true, solution: false -->
```python
def test_correlation_title_length_salary():
    """Check that correlation_title_length_salary returns the correct output for several inputs"""

# print("Testing correlation_title_length_salary()...", correlation_title_length_salary(df))
```

Output:
```
Testing correlation_title_length_salary()... -0.03687844593260675
```